% Refutation found. Thanks to Tanya!
% SZS status Theorem for SEU002+1
% SZS output start Proof for SEU002+1
fof(f4,axiom,(
  ! [X0] : ((function(X0) & relation(X0)) => ! [X1] : (relation_rng(X0) = X1 <=> ! [X2] : (in(X2,X1) <=> ? [X3] : (apply(X0,X3) = X2 & in(X3,relation_dom(X0))))))),
  file('Problems/SEU/SEU002+1.p',unknown)).
fof(f5,axiom,(
  ! [X0,X1] : ((relation(X1) & relation(X0)) => relation(relation_composition(X0,X1)))),
  file('Problems/SEU/SEU002+1.p',unknown)).
fof(f9,axiom,(
  ! [X0,X1] : ((function(X1) & relation(X1) & function(X0) & relation(X0)) => (function(relation_composition(X0,X1)) & relation(relation_composition(X0,X1))))),
  file('Problems/SEU/SEU002+1.p',unknown)).
fof(f28,axiom,(
  ! [X0,X1] : ((function(X1) & relation(X1)) => ! [X2] : ((function(X2) & relation(X2)) => (in(X0,relation_dom(relation_composition(X2,X1))) <=> (in(apply(X2,X0),relation_dom(X1)) & in(X0,relation_dom(X2))))))),
  file('Problems/SEU/SEU002+1.p',unknown)).
fof(f29,axiom,(
  ! [X0,X1] : ((function(X1) & relation(X1)) => ! [X2] : ((function(X2) & relation(X2)) => (in(X0,relation_dom(relation_composition(X2,X1))) => apply(X1,apply(X2,X0)) = apply(relation_composition(X2,X1),X0))))),
  file('Problems/SEU/SEU002+1.p',unknown)).
fof(f30,conjecture,(
  ! [X0,X1] : ((function(X1) & relation(X1)) => ! [X2] : ((function(X2) & relation(X2)) => (in(X0,relation_rng(relation_composition(X2,X1))) => in(X0,relation_rng(X1)))))),
  file('Problems/SEU/SEU002+1.p',unknown)).
fof(f31,negated_conjecture,(
  ~! [X0,X1] : ((function(X1) & relation(X1)) => ! [X2] : ((function(X2) & relation(X2)) => (in(X0,relation_rng(relation_composition(X2,X1))) => in(X0,relation_rng(X1)))))),
  inference(negated_conjecture,[],[f30])).
fof(f46,plain,(
  ! [X0] : (! [X1] : (relation_rng(X0) = X1 <=> ! [X2] : (in(X2,X1) <=> ? [X3] : (apply(X0,X3) = X2 & in(X3,relation_dom(X0))))) | (~function(X0) | ~relation(X0)))),
  inference(ennf_transformation,[],[f4])).
fof(f47,plain,(
  ! [X0] : (! [X1] : (relation_rng(X0) = X1 <=> ! [X2] : (in(X2,X1) <=> ? [X3] : (apply(X0,X3) = X2 & in(X3,relation_dom(X0))))) | ~function(X0) | ~relation(X0))),
  inference(flattening,[],[f46])).
fof(f48,plain,(
  ! [X0,X1] : (relation(relation_composition(X0,X1)) | (~relation(X1) | ~relation(X0)))),
  inference(ennf_transformation,[],[f5])).
fof(f49,plain,(
  ! [X0,X1] : (relation(relation_composition(X0,X1)) | ~relation(X1) | ~relation(X0))),
  inference(flattening,[],[f48])).
fof(f52,plain,(
  ! [X0,X1] : ((function(relation_composition(X0,X1)) & relation(relation_composition(X0,X1))) | (~function(X1) | ~relation(X1) | ~function(X0) | ~relation(X0)))),
  inference(ennf_transformation,[],[f9])).
fof(f53,plain,(
  ! [X0,X1] : ((function(relation_composition(X0,X1)) & relation(relation_composition(X0,X1))) | ~function(X1) | ~relation(X1) | ~function(X0) | ~relation(X0))),
  inference(flattening,[],[f52])).
fof(f64,plain,(
  ! [X0,X1] : (! [X2] : ((in(X0,relation_dom(relation_composition(X2,X1))) <=> (in(apply(X2,X0),relation_dom(X1)) & in(X0,relation_dom(X2)))) | (~function(X2) | ~relation(X2))) | (~function(X1) | ~relation(X1)))),
  inference(ennf_transformation,[],[f28])).
fof(f65,plain,(
  ! [X0,X1] : (! [X2] : ((in(X0,relation_dom(relation_composition(X2,X1))) <=> (in(apply(X2,X0),relation_dom(X1)) & in(X0,relation_dom(X2)))) | ~function(X2) | ~relation(X2)) | ~function(X1) | ~relation(X1))),
  inference(flattening,[],[f64])).
fof(f66,plain,(
  ! [X0,X1] : (! [X2] : ((apply(X1,apply(X2,X0)) = apply(relation_composition(X2,X1),X0) | ~in(X0,relation_dom(relation_composition(X2,X1)))) | (~function(X2) | ~relation(X2))) | (~function(X1) | ~relation(X1)))),
  inference(ennf_transformation,[],[f29])).
fof(f67,plain,(
  ! [X0,X1] : (! [X2] : (apply(X1,apply(X2,X0)) = apply(relation_composition(X2,X1),X0) | ~in(X0,relation_dom(relation_composition(X2,X1))) | ~function(X2) | ~relation(X2)) | ~function(X1) | ~relation(X1))),
  inference(flattening,[],[f66])).
fof(f68,plain,(
  ? [X0,X1] : (? [X2] : ((~in(X0,relation_rng(X1)) & in(X0,relation_rng(relation_composition(X2,X1)))) & (function(X2) & relation(X2))) & (function(X1) & relation(X1)))),
  inference(ennf_transformation,[],[f31])).
fof(f69,plain,(
  ? [X0,X1] : (? [X2] : (~in(X0,relation_rng(X1)) & in(X0,relation_rng(relation_composition(X2,X1))) & function(X2) & relation(X2)) & function(X1) & relation(X1))),
  inference(flattening,[],[f68])).
fof(f79,plain,(
  ! [X0] : (! [X1] : ((relation_rng(X0) = X1 | ? [X2] : ((! [X3] : (apply(X0,X3) \= X2 | ~in(X3,relation_dom(X0))) | ~in(X2,X1)) & (? [X3] : (apply(X0,X3) = X2 & in(X3,relation_dom(X0))) | in(X2,X1)))) & (! [X2] : ((in(X2,X1) | ! [X3] : (apply(X0,X3) \= X2 | ~in(X3,relation_dom(X0)))) & (? [X3] : (apply(X0,X3) = X2 & in(X3,relation_dom(X0))) | ~in(X2,X1))) | relation_rng(X0) \= X1)) | ~function(X0) | ~relation(X0))),
  inference(nnf_transformation,[],[f47])).
fof(f80,plain,(
  ! [X0] : (! [X1] : ((relation_rng(X0) = X1 | ? [X2] : ((! [X3] : (apply(X0,X3) \= X2 | ~in(X3,relation_dom(X0))) | ~in(X2,X1)) & (? [X4] : (apply(X0,X4) = X2 & in(X4,relation_dom(X0))) | in(X2,X1)))) & (! [X5] : ((in(X5,X1) | ! [X6] : (apply(X0,X6) \= X5 | ~in(X6,relation_dom(X0)))) & (? [X7] : (apply(X0,X7) = X5 & in(X7,relation_dom(X0))) | ~in(X5,X1))) | relation_rng(X0) \= X1)) | ~function(X0) | ~relation(X0))),
  inference(rectify,[],[f79])).
fof(f81,plain,(
  ! [X1,X0] : (? [X2] : ((! [X3] : (apply(X0,X3) \= X2 | ~in(X3,relation_dom(X0))) | ~in(X2,X1)) & (? [X4] : (apply(X0,X4) = X2 & in(X4,relation_dom(X0))) | in(X2,X1))) => ((! [X3] : (apply(X0,X3) \= sK0(X0,X1) | ~in(X3,relation_dom(X0))) | ~in(sK0(X0,X1),X1)) & (? [X4] : (apply(X0,X4) = sK0(X0,X1) & in(X4,relation_dom(X0))) | in(sK0(X0,X1),X1))))),
  introduced(choice_axiom,[])).
fof(f82,plain,(
  ( ! [X2] : (! [X1,X0] : (? [X4] : (apply(X0,X4) = X2 & in(X4,relation_dom(X0))) => (apply(X0,sK1(X0,X1)) = X2 & in(sK1(X0,X1),relation_dom(X0))))) )),
  introduced(choice_axiom,[])).
fof(f83,plain,(
  ! [X5,X0] : (? [X7] : (apply(X0,X7) = X5 & in(X7,relation_dom(X0))) => (apply(X0,sK2(X0,X5)) = X5 & in(sK2(X0,X5),relation_dom(X0))))),
  introduced(choice_axiom,[])).
fof(f84,plain,(
  ! [X0] : (! [X1] : ((relation_rng(X0) = X1 | ((! [X3] : (apply(X0,X3) \= sK0(X0,X1) | ~in(X3,relation_dom(X0))) | ~in(sK0(X0,X1),X1)) & ((apply(X0,sK1(X0,X1)) = sK0(X0,X1) & in(sK1(X0,X1),relation_dom(X0))) | in(sK0(X0,X1),X1)))) & (! [X5] : ((in(X5,X1) | ! [X6] : (apply(X0,X6) \= X5 | ~in(X6,relation_dom(X0)))) & ((apply(X0,sK2(X0,X5)) = X5 & in(sK2(X0,X5),relation_dom(X0))) | ~in(X5,X1))) | relation_rng(X0) \= X1)) | ~function(X0) | ~relation(X0))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f80,f83,f82,f81])).
fof(f103,plain,(
  ! [X0,X1] : (! [X2] : (((in(X0,relation_dom(relation_composition(X2,X1))) | (~in(apply(X2,X0),relation_dom(X1)) | ~in(X0,relation_dom(X2)))) & ((in(apply(X2,X0),relation_dom(X1)) & in(X0,relation_dom(X2))) | ~in(X0,relation_dom(relation_composition(X2,X1))))) | ~function(X2) | ~relation(X2)) | ~function(X1) | ~relation(X1))),
  inference(nnf_transformation,[],[f65])).
fof(f104,plain,(
  ! [X0,X1] : (! [X2] : (((in(X0,relation_dom(relation_composition(X2,X1))) | ~in(apply(X2,X0),relation_dom(X1)) | ~in(X0,relation_dom(X2))) & ((in(apply(X2,X0),relation_dom(X1)) & in(X0,relation_dom(X2))) | ~in(X0,relation_dom(relation_composition(X2,X1))))) | ~function(X2) | ~relation(X2)) | ~function(X1) | ~relation(X1))),
  inference(flattening,[],[f103])).
fof(f105,plain,(
  ? [X0,X1] : (? [X2] : (~in(X0,relation_rng(X1)) & in(X0,relation_rng(relation_composition(X2,X1))) & function(X2) & relation(X2)) & function(X1) & relation(X1)) => (? [X2] : (~in(sK12,relation_rng(sK13)) & in(sK12,relation_rng(relation_composition(X2,sK13))) & function(X2) & relation(X2)) & function(sK13) & relation(sK13))),
  introduced(choice_axiom,[])).
fof(f106,plain,(
  ( ! [X0,X1] : (? [X2] : (~in(X0,relation_rng(X1)) & in(X0,relation_rng(relation_composition(X2,X1))) & function(X2) & relation(X2)) => (~in(X0,relation_rng(X1)) & in(X0,relation_rng(relation_composition(sK14,X1))) & function(sK14) & relation(sK14))) )),
  introduced(choice_axiom,[])).
fof(f107,plain,(
  (~in(sK12,relation_rng(sK13)) & in(sK12,relation_rng(relation_composition(sK14,sK13))) & function(sK14) & relation(sK14)) & function(sK13) & relation(sK13)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13,sK14])],[f69,f106,f105])).
fof(f111,plain,(
  ( ! [X0,X5,X1] : (in(sK2(X0,X5),relation_dom(X0)) | ~in(X5,X1) | relation_rng(X0) \= X1 | ~function(X0) | ~relation(X0)) )),
  inference(cnf_transformation,[],[f84])).
fof(f112,plain,(
  ( ! [X0,X5,X1] : (apply(X0,sK2(X0,X5)) = X5 | ~in(X5,X1) | relation_rng(X0) \= X1 | ~function(X0) | ~relation(X0)) )),
  inference(cnf_transformation,[],[f84])).
fof(f113,plain,(
  ( ! [X6,X0,X5,X1] : (in(X5,X1) | apply(X0,X6) \= X5 | ~in(X6,relation_dom(X0)) | relation_rng(X0) \= X1 | ~function(X0) | ~relation(X0)) )),
  inference(cnf_transformation,[],[f84])).
fof(f117,plain,(
  ( ! [X0,X1] : (relation(relation_composition(X0,X1)) | ~relation(X1) | ~relation(X0)) )),
  inference(cnf_transformation,[],[f49])).
fof(f124,plain,(
  ( ! [X0,X1] : (function(relation_composition(X0,X1)) | ~function(X1) | ~relation(X1) | ~function(X0) | ~relation(X0)) )),
  inference(cnf_transformation,[],[f53])).
fof(f153,plain,(
  ( ! [X2,X0,X1] : (~in(X0,relation_dom(relation_composition(X2,X1))) | in(apply(X2,X0),relation_dom(X1)) | ~function(X2) | ~relation(X2) | ~function(X1) | ~relation(X1)) )),
  inference(cnf_transformation,[],[f104])).
fof(f155,plain,(
  ( ! [X2,X0,X1] : (~in(X0,relation_dom(relation_composition(X2,X1))) | apply(X1,apply(X2,X0)) = apply(relation_composition(X2,X1),X0) | ~function(X2) | ~relation(X2) | ~function(X1) | ~relation(X1)) )),
  inference(cnf_transformation,[],[f67])).
fof(f156,plain,(
  relation(sK13)),
  inference(cnf_transformation,[],[f107])).
fof(f157,plain,(
  function(sK13)),
  inference(cnf_transformation,[],[f107])).
fof(f158,plain,(
  relation(sK14)),
  inference(cnf_transformation,[],[f107])).
fof(f159,plain,(
  function(sK14)),
  inference(cnf_transformation,[],[f107])).
fof(f160,plain,(
  in(sK12,relation_rng(relation_composition(sK14,sK13)))),
  inference(cnf_transformation,[],[f107])).
fof(f161,plain,(
  ~in(sK12,relation_rng(sK13))),
  inference(cnf_transformation,[],[f107])).
fof(f169,plain,(
  ( ! [X6,X0,X1] : (in(apply(X0,X6),X1) | ~in(X6,relation_dom(X0)) | relation_rng(X0) \= X1 | ~function(X0) | ~relation(X0)) )),
  inference(equality_resolution,[],[f113])).
fof(f170,plain,(
  ( ! [X6,X0] : (in(apply(X0,X6),relation_rng(X0)) | ~in(X6,relation_dom(X0)) | ~function(X0) | ~relation(X0)) )),
  inference(equality_resolution,[],[f169])).
fof(f171,plain,(
  ( ! [X0,X5] : (~in(X5,relation_rng(X0)) | apply(X0,sK2(X0,X5)) = X5 | ~function(X0) | ~relation(X0)) )),
  inference(equality_resolution,[],[f112])).
fof(f172,plain,(
  ( ! [X0,X5] : (in(sK2(X0,X5),relation_dom(X0)) | ~in(X5,relation_rng(X0)) | ~function(X0) | ~relation(X0)) )),
  inference(equality_resolution,[],[f111])).
fof(f294,plain,(
  apply(relation_composition(sK14,sK13),sK2(relation_composition(sK14,sK13),sK12)) = sK12 | ~function(relation_composition(sK14,sK13)) | ~relation(relation_composition(sK14,sK13))),
  inference(resolution,[],[f171,f160])).
fof(f298,plain,(
  spl15_4 <=> relation(relation_composition(sK14,sK13))),
  introduced(avatar_definition,[new_symbols(naming,[spl15_4])])).
fof(f299,plain,(
  relation(relation_composition(sK14,sK13)) | ~spl15_4),
  inference(avatar_component_clause,[],[f298])).
fof(f301,plain,(
  spl15_5 <=> ~relation(relation_composition(sK14,sK13))),
  introduced(avatar_definition,[new_symbols(naming,[spl15_5])])).
fof(f302,plain,(
  ~relation(relation_composition(sK14,sK13)) | ~spl15_5),
  inference(avatar_component_clause,[],[f301])).
fof(f304,plain,(
  spl15_6 <=> function(relation_composition(sK14,sK13))),
  introduced(avatar_definition,[new_symbols(naming,[spl15_6])])).
fof(f305,plain,(
  function(relation_composition(sK14,sK13)) | ~spl15_6),
  inference(avatar_component_clause,[],[f304])).
fof(f307,plain,(
  spl15_7 <=> ~function(relation_composition(sK14,sK13))),
  introduced(avatar_definition,[new_symbols(naming,[spl15_7])])).
fof(f308,plain,(
  ~function(relation_composition(sK14,sK13)) | ~spl15_7),
  inference(avatar_component_clause,[],[f307])).
fof(f313,plain,(
  spl15_8 <=> apply(relation_composition(sK14,sK13),sK2(relation_composition(sK14,sK13),sK12)) = sK12),
  introduced(avatar_definition,[new_symbols(naming,[spl15_8])])).
fof(f314,plain,(
  apply(relation_composition(sK14,sK13),sK2(relation_composition(sK14,sK13),sK12)) = sK12 | ~spl15_8),
  inference(avatar_component_clause,[],[f313])).
fof(f315,plain,(
  ~spl15_5 | ~spl15_7 | spl15_8),
  inference(avatar_split_clause,[],[f294,f313,f307,f301])).
fof(f318,plain,(
  ~relation(sK13) | ~relation(sK14) | ~spl15_5),
  inference(resolution,[],[f302,f117])).
fof(f319,plain,(
  ~relation(sK14) | ~spl15_5),
  inference(subsumption_resolution,[],[f318,f156])).
fof(f320,plain,(
  $false | ~spl15_5),
  inference(subsumption_resolution,[],[f319,f158])).
fof(f321,plain,(
  spl15_5),
  inference(avatar_contradiction_clause,[],[f320])).
fof(f343,plain,(
  ~function(sK13) | ~relation(sK13) | ~function(sK14) | ~relation(sK14) | ~spl15_7),
  inference(resolution,[],[f308,f124])).
fof(f344,plain,(
  ~relation(sK13) | ~function(sK14) | ~relation(sK14) | ~spl15_7),
  inference(subsumption_resolution,[],[f343,f157])).
fof(f345,plain,(
  ~function(sK14) | ~relation(sK14) | ~spl15_7),
  inference(subsumption_resolution,[],[f344,f156])).
fof(f346,plain,(
  ~relation(sK14) | ~spl15_7),
  inference(subsumption_resolution,[],[f345,f159])).
fof(f347,plain,(
  $false | ~spl15_7),
  inference(subsumption_resolution,[],[f346,f158])).
fof(f348,plain,(
  spl15_7),
  inference(avatar_contradiction_clause,[],[f347])).
fof(f377,plain,(
  ( ! [X10,X11,X9] : (apply(X9,apply(X10,sK2(relation_composition(X10,X9),X11))) = apply(relation_composition(X10,X9),sK2(relation_composition(X10,X9),X11)) | ~function(X10) | ~relation(X10) | ~function(X9) | ~relation(X9) | ~in(X11,relation_rng(relation_composition(X10,X9))) | ~function(relation_composition(X10,X9)) | ~relation(relation_composition(X10,X9))) )),
  inference(resolution,[],[f155,f172])).
fof(f381,plain,(
  ( ! [X10,X11,X9] : (apply(X9,apply(X10,sK2(relation_composition(X10,X9),X11))) = apply(relation_composition(X10,X9),sK2(relation_composition(X10,X9),X11)) | ~function(X10) | ~relation(X10) | ~function(X9) | ~relation(X9) | ~in(X11,relation_rng(relation_composition(X10,X9))) | ~relation(relation_composition(X10,X9))) )),
  inference(subsumption_resolution,[],[f377,f124])).
fof(f382,plain,(
  ( ! [X10,X11,X9] : (~in(X11,relation_rng(relation_composition(X10,X9))) | ~function(X10) | ~relation(X10) | ~function(X9) | ~relation(X9) | apply(X9,apply(X10,sK2(relation_composition(X10,X9),X11))) = apply(relation_composition(X10,X9),sK2(relation_composition(X10,X9),X11))) )),
  inference(subsumption_resolution,[],[f381,f117])).
fof(f536,plain,(
  spl15_10 <=> in(sK2(relation_composition(sK14,sK13),sK12),relation_dom(relation_composition(sK14,sK13)))),
  introduced(avatar_definition,[new_symbols(naming,[spl15_10])])).
fof(f537,plain,(
  in(sK2(relation_composition(sK14,sK13),sK12),relation_dom(relation_composition(sK14,sK13))) | ~spl15_10),
  inference(avatar_component_clause,[],[f536])).
fof(f539,plain,(
  spl15_11 <=> ~in(sK2(relation_composition(sK14,sK13),sK12),relation_dom(relation_composition(sK14,sK13)))),
  introduced(avatar_definition,[new_symbols(naming,[spl15_11])])).
fof(f540,plain,(
  ~in(sK2(relation_composition(sK14,sK13),sK12),relation_dom(relation_composition(sK14,sK13))) | ~spl15_11),
  inference(avatar_component_clause,[],[f539])).
fof(f1450,plain,(
  ~function(sK14) | ~relation(sK14) | ~function(sK13) | ~relation(sK13) | apply(relation_composition(sK14,sK13),sK2(relation_composition(sK14,sK13),sK12)) = apply(sK13,apply(sK14,sK2(relation_composition(sK14,sK13),sK12)))),
  inference(resolution,[],[f382,f160])).
fof(f1465,plain,(
  ~relation(sK14) | ~function(sK13) | ~relation(sK13) | apply(relation_composition(sK14,sK13),sK2(relation_composition(sK14,sK13),sK12)) = apply(sK13,apply(sK14,sK2(relation_composition(sK14,sK13),sK12)))),
  inference(subsumption_resolution,[],[f1450,f159])).
fof(f1466,plain,(
  ~function(sK13) | ~relation(sK13) | apply(relation_composition(sK14,sK13),sK2(relation_composition(sK14,sK13),sK12)) = apply(sK13,apply(sK14,sK2(relation_composition(sK14,sK13),sK12)))),
  inference(subsumption_resolution,[],[f1465,f158])).
fof(f1467,plain,(
  ~relation(sK13) | apply(relation_composition(sK14,sK13),sK2(relation_composition(sK14,sK13),sK12)) = apply(sK13,apply(sK14,sK2(relation_composition(sK14,sK13),sK12)))),
  inference(subsumption_resolution,[],[f1466,f157])).
fof(f1468,plain,(
  apply(relation_composition(sK14,sK13),sK2(relation_composition(sK14,sK13),sK12)) = apply(sK13,apply(sK14,sK2(relation_composition(sK14,sK13),sK12)))),
  inference(subsumption_resolution,[],[f1467,f156])).
fof(f1469,plain,(
  apply(sK13,apply(sK14,sK2(relation_composition(sK14,sK13),sK12))) = sK12 | ~spl15_8),
  inference(forward_demodulation,[],[f1468,f314])).
fof(f2704,plain,(
  ~in(sK12,relation_rng(relation_composition(sK14,sK13))) | ~function(relation_composition(sK14,sK13)) | ~relation(relation_composition(sK14,sK13)) | ~spl15_11),
  inference(resolution,[],[f540,f172])).
fof(f2705,plain,(
  ~function(relation_composition(sK14,sK13)) | ~relation(relation_composition(sK14,sK13)) | ~spl15_11),
  inference(subsumption_resolution,[],[f2704,f160])).
fof(f2706,plain,(
  ~relation(relation_composition(sK14,sK13)) | (~spl15_6 | ~spl15_11)),
  inference(subsumption_resolution,[],[f2705,f305])).
fof(f2707,plain,(
  $false | (~spl15_4 | ~spl15_6 | ~spl15_11)),
  inference(subsumption_resolution,[],[f2706,f299])).
fof(f2708,plain,(
  ~spl15_4 | ~spl15_6 | spl15_11),
  inference(avatar_contradiction_clause,[],[f2707])).
fof(f2710,plain,(
  in(apply(sK14,sK2(relation_composition(sK14,sK13),sK12)),relation_dom(sK13)) | ~function(sK14) | ~relation(sK14) | ~function(sK13) | ~relation(sK13) | ~spl15_10),
  inference(resolution,[],[f537,f153])).
fof(f2722,plain,(
  in(apply(sK14,sK2(relation_composition(sK14,sK13),sK12)),relation_dom(sK13)) | ~relation(sK14) | ~function(sK13) | ~relation(sK13) | ~spl15_10),
  inference(subsumption_resolution,[],[f2710,f159])).
fof(f2723,plain,(
  in(apply(sK14,sK2(relation_composition(sK14,sK13),sK12)),relation_dom(sK13)) | ~function(sK13) | ~relation(sK13) | ~spl15_10),
  inference(subsumption_resolution,[],[f2722,f158])).
fof(f2724,plain,(
  in(apply(sK14,sK2(relation_composition(sK14,sK13),sK12)),relation_dom(sK13)) | ~relation(sK13) | ~spl15_10),
  inference(subsumption_resolution,[],[f2723,f157])).
fof(f2725,plain,(
  in(apply(sK14,sK2(relation_composition(sK14,sK13),sK12)),relation_dom(sK13)) | ~spl15_10),
  inference(subsumption_resolution,[],[f2724,f156])).
fof(f2926,plain,(
  in(sK12,relation_rng(sK13)) | ~in(apply(sK14,sK2(relation_composition(sK14,sK13),sK12)),relation_dom(sK13)) | ~function(sK13) | ~relation(sK13) | ~spl15_8),
  inference(superposition,[],[f170,f1469])).
fof(f2939,plain,(
  ~in(apply(sK14,sK2(relation_composition(sK14,sK13),sK12)),relation_dom(sK13)) | ~function(sK13) | ~relation(sK13) | ~spl15_8),
  inference(subsumption_resolution,[],[f2926,f161])).
fof(f2940,plain,(
  ~function(sK13) | ~relation(sK13) | (~spl15_8 | ~spl15_10)),
  inference(subsumption_resolution,[],[f2939,f2725])).
fof(f2941,plain,(
  ~relation(sK13) | (~spl15_8 | ~spl15_10)),
  inference(subsumption_resolution,[],[f2940,f157])).
fof(f2942,plain,(
  $false | (~spl15_8 | ~spl15_10)),
  inference(subsumption_resolution,[],[f2941,f156])).
fof(f2943,plain,(
  ~spl15_8 | ~spl15_10),
  inference(avatar_contradiction_clause,[],[f2942])).
fof(f2944,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f315,f321,f348,f2708,f2943])).
% SZS output end Proof for SEU002+1
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 6524
% Time elapsed: 0.090 s
% ------------------------------
% ------------------------------
